# Homework 3: More String, Lists, and Python Objects (18 pts)

name: Bryant Willoughby     

email: btwillou@umich.edu

This homework assignment took me XX hours in total to complete. (Please help us to gauge the difficulty of the assignment.)

## Collaboration Disclosure

In the cell below, please list *everyone* with whom you discussed any of the homework problems, excluding only the GSIs and the course instructor. 

If you did not discuss the homework with anyone else, write __"I did not discuss this homework with anyone."__

Even if you discuss questions with other, the code you submit must be only yours. All work is checked with the [MOSS plagiarism detector](https://theory.stanford.edu/~aiken/moss/).

__"I did not discuss this homework with anyone."__

---

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. 


Before submitting, please make sure to __Cells->Run All__ executes without errors; errors in your code translate directly to point deductions. 
In general, you don't need to do explicitly raise errors (e.g. with the ```raise``` function) if we don't ask you to in the problem statement.
However, even in cases where we ask you to check for errors, your submission should not contain any examples of your functions actually raising those errors.

Note that many parts of this homework where you are expected to type in code will have ```NotImplementedError()``` as a placeholder. You need to delete this function and replace it with your own code.

## Importing libaries

Unless specifically prohibited, you can import any Python library that comes included with Python (e.g. ```math```, ```datetime```) or that has been used in the lecture or homework prompts. You do not need any other libraries to complete the homework. Note that the graders will not install any additional libraries to grade an assignment.


<!-- ## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking: -->

## Homework tips 

1. **Start early!** If you run into trouble installing things or importing packages, it’s
best to find those problems well in advance, not the night before your assignment is
due when we cannot help you!

2. **Make sure you back up your work!** At a minimum, do your work in a Dropbox
folder. Better yet, use git, which is well worth your time and effort to learn.

3. **Be careful to follow directions!** Remember that Python is case sensitive. If
you are ask you to define a function called my_function and you define a function
called My_Function, you will not receive full credit. You may want to copy-paste
the function names below to make sure that the functions in your notebook match.

## Error checking

You do not need to do error checking (raising errors, etc.) in your code unless we explicitly ask you to so in a problem.

## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking:


## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. You will notice some `read-only` cells in the assignment that contain `assert` statements. These are tests that your code must pass for your solution to be correct. If any of the tests fail, you will get an python error and not get points for that question. 

**Note:** The tests shown not are not comprehensive; additional tests will be used at grading time. You are encouraged to read the problem carefully and verify your code covers all possible cases.

**Be careful:** If a jupyter notebook cell takes longer than `60s` to run, the autograder will not grade it and you will receive zero credit for that question.

```python
import collections
import string
import itertools
```

---
## Question 1: Infinite sequences (8 pts)
For each of the problems below, write a function which generates the given infinite sequence. We should be able to use your generators to access any entry of the sequence no matter how deep.

**1(a)** (2 pts) The prime numbers are

$$2, 3, 5, 7, 11, 13, 17, \dots$$

Give a generator for the primes. (Note: many algorithms exist for this problem. Yours should be efficient enough that we can use it to generate reasonably large prime numbers.)

The generator yields primes by testing each integer $n \gt 2$ and remembering primes found so far. For each candidate n, it tests divisibility only by those previous primes p and stops the inner loop once $p^2 \gt n$. This works because any nontrivial factor of n must be $\le \sqrt(n)$; if no divisor is found, n is prime, append it to the list, and yield. The function is an infinite generator (it never raises `StopIteration` itself). 

<!-- grade_id: primes, points: 0, grade: false, solution: true -->
```python
def primes():
    """Yield the prime numbers indefinitely: 2, 3, 5, 7, ...

    Implementation notes:
    - We remember primes found so far and test new candidates by trial
      division using only those primes, stopping when p*p > n.
    """
    primes_list = []  # list of primes discovered so far (used as trial divisors)
    n = 2              # current integer to test for primality (start at first prime)
    while True:
        is_prime = True  # assume n is prime until a divisor is found
        # Test divisibility by known primes; only need to test up to sqrt(n).
        for p in primes_list:
            # If p*p > n then no further primes can divide n -> stop checks early.
            if p * p > n:
                break
            # If n is divisible by any known prime p, it's composite.
            if n % p == 0:
                is_prime = False
                break
        if is_prime:
            # Record the new prime so future checks can use it, then yield it.
            primes_list.append(n)
            yield n
        n += 1  # move to next integer and repeat (generator is infinite)

# p = primes()
# plist = [2,3,5,7,11,13,17,19,23,29]
# for i in range(10):
#     assert next(p) == plist[i]
# print(plist)
```

<!-- grade_id: correct_prime, points: 2, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
p = primes()
plist = [2,3,5,7,11,13,17,19,23,29]
for i in range(10):
    assert next(p) == plist[i]
```

**1(b)** (3 pts) The *ruler sequence* is

$$1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7\dots$$

The first number (1) appears once; the next two numbers (2 and 3) appear twice, the next three numbers appear three times, etc.

Hint: You are highly encouraged to use relevant functions from itertools to make your solution simple

This generator yields blocks of increasing size: for k = 1, 2, 3... and outputs k distinct consecutive integers and repeats each of those integers k times. This implementation uses `itertools.count` and `itertools.repeat` so that the repetition logic is expressed clearly and efficiently without nested manual loops. 

```python
# help(itertools.count); help(itertools.repeat)

def ruler():
    """Ruler-like sequence using itertools for concise repetition.

    Uses itertools.count to iterate block sizes 1,2,3,... and itertools.repeat
    to emit each integer the required number of times.
    """
    current = 1
    # starts count at 1 and counts up indefinitely
    for block_size in itertools.count(1):
        # produce `block_size` distinct integers in this block
        for _ in range(block_size):
            # yield the same `current` value exactly `block_size` times
            yield from itertools.repeat(current, block_size)
            # advance to the next integer after finishing repeats for `current`
            current += 1

# r = ruler()
# ruler_numbers = [1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7, 8, 8, 8, 8]
# for i in range(22):
#     assert next(r) == ruler_numbers[i]
# print(ruler_numbers)
```

<!-- grade_id: correct_ruler, points: 3, grade: true, solution: false -->
```python
r = ruler()
ruler_numbers = [1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7, 8, 8, 8, 8]
for n in ruler_numbers:
    assert next(r) == n
```

**1(c)** (3 pts) The look-and-say sequence  

$$1, 11, 21, 1211, 111221, 312211, 13112221, 1113213211, \dots$$

is generated as follows:

1. The first entry is one.
2. The second entry is generated by reading the first entry aloud: "one one"
3. The third entry is generated by reading the second entry aloud: "two ones"
4. The fourth entry is generated by reading the third entry aloud: "one two and one one"
5. The fifth entry is therefore "one one, one two, and two ones"

... and so forth. Note that each entry of the sequence should be a `str` object.

This generator produces the look-and-say sequence by repeatedly "reading" the current string s: it counts consecutive identical characters and writes each run as the count followed by the character to form the next term. Each call to `next()` returns the current term (a string) and the generator then computes the following term by scanning runs and concatenating count \& digit pairs. 

<!-- grade_id: look_say, points: 0, grade: false, solution: true -->
```python
# help(str.join)
def look_say():
    """Generator for the look-and-say sequence.

    Each term is produced by reading off the digits of the previous term
    counting consecutive runs. The sequence starts at '1'. Every yielded
    value is a string.
    """
    s = '1'                 # current term as a string (start value)
    while True:
        # yield the current term to the caller; generator pauses here
        yield s

        # Build the next term by scanning 's' and counting runs of identical chars.
        parts = []         # collect components (count, digit) as strings
        i = 0              # read-head index into s

        # Process the whole string s by advancing i to the end of each run.
        while i < len(s):
            j = i + 1
            # advance j while characters match the current run
            while j < len(s) and s[j] == s[i]:
                j += 1
            # j is now one past the end of the run starting at i
            count = j - i               # length of the run (number of repeated chars)
            parts.append(str(count))    # append the run length as text
            parts.append(s[i])          # append the digit/character being counted
            i = j                       # continue scanning from the end of this run

        # Join the collected parts into the next term string, e.g. '111221'
        s = ''.join(parts)

# s = look_say()
# look_say_first_ten = ['1',
#  '11',
#  '21',
#  '1211',
#  '111221',
#  '312211',
#  '13112221',
#  '1113213211',
#  '31131211131221',
#  '13211311123113112211']
# for i in range(10):
#     assert next(s) == look_say_first_ten[i]
# print(look_say_first_ten)
```

<!-- grade_id: correct_look_say, points: 3, grade: true, solution: false -->
```python
s = look_say()
look_say_first_ten = ['1',
 '11',
 '21',
 '1211',
 '111221',
 '312211',
 '13112221',
 '1113213211',
 '31131211131221',
 '13211311123113112211']
for i in range(10):
    s_i = next(s)
    assert isinstance(s_i, str)
    assert s_i == look_say_first_ten[i]
```

### Question 2: Simple ciphers (5 pts)

A *cipher* is an algorithm for encrypting or decrypting a text message, called *plaintext*, into *ciphertext*. In this exercise, we will assume that all plaintext to be encrypted or decrypted are composed of the lower-case Roman alphabet, without any punctuation or whitespace. Examples of such messages could be

```attackatdawn```

or 

```iamajellyfilleddonut```.

This function implements a Caesar cipher: each lower-case letter in s is converted to a number 0-25, shifted by k (with wrap-around), and converted back to a letter. k is reduced modulo 26 so positive or negative shifts behave correctly; the function returns the encrypted string. 

<!-- grade_id: enc_caesar, points: 0, grade: false, solution: true -->
```python
# help(ord) 

def enc_caesar(s: str, k: int) -> str:
    """
    Encrypt the message s using a Caesar cipher with key k.

    Assumptions: s contains only lower-case letters a-z (no spaces or punctuation).
    Letters are shifted forward by k (wrapping modulo 26). k may be negative.
    """
    # Normalize shift to range 0..25 so large or negative k values wrap correctly
    k = k % 26

    out_chars = []                # collect shifted characters here
    for ch in s:
        # Convert 'a'..'z' to numbers 0..25
        o = ord(ch) - ord('a')
        # Apply the shift, wrap with %26, then convert back to a character
        out_chars.append(chr((o + k) % 26 + ord('a')))
    # Join the list of characters into the resulting encrypted string
    return ''.join(out_chars)

# public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
# for msg, k, enc in public_tests:
#     assert enc_caesar(msg, k) == enc
# print(public_tests)
```

<!-- grade_id: correct_enc_caesar, points: 2, grade: true, solution: false -->
```python
public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
for msg, k, enc in public_tests:
    assert enc_caesar(msg, k) == enc
```

<!-- grade_id: dec_caesar, points: 0, grade: false, solution: true -->
```python
def dec_caesar(s: str, k: int) -> str:
    """
    Decrypt the message s encrypted with a Caesar cipher with key k.

    This simply applies the negative shift of enc_caesar.
    """
    # decrypt by shifting backward by k (i.e., forward by -k)
    return enc_caesar(s, -k)

# public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
# for msg, k, enc in public_tests: 
#     assert dec_caesar(enc, k) == msg
# public_tests
```

<!-- grade_id: correct_dec_caesar, points: 3, grade: true, solution: false -->
```python
public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
for msg, k, enc in public_tests: 
    assert dec_caesar(enc, k) == msg
```

### Question 3 (5 pts) Valid String.
Given a string s containing just the characters '(', ')', '{', '}', '[' and ']', determine if the input string is valid.

An input string is valid if:

1) Open brackets must be closed by the same type of brackets.
2) Open brackets must be closed in the correct order.
3) Every close bracket has a corresponding open bracket of the same type.

This function checks whether a string of brackets is properly balanced and nested by using a stack: it pushes opening brackets, and for each closing bracket, it pops the most recent opener and verifies it matches the correct type. It returns `False` immediately on any mismatch or if a closer lacks an opener, and returns `True` only if all openers have been matched (stack is empty) at the end. 

```python
def is_valid(s):
    """Check bracket validity using a stack.
    
    Args:
      s (str): string containing only the characters '()', '[]', '{}'.
    Returns:
      bool: True if every closer matches the most recent opener and all openers are closed.
    """
    stack = []  # holds unmatched opening brackets
    for char in s:
        # push opening brackets
        if char in '([{':
            stack.append(char)
        # on closing bracket, there must be a matching opener on top of the stack
        elif char in ')]}':
            if not stack:           # no opener available to match
                return False
            last = stack.pop()      # most recent unmatched opener
            # ensure types match: e.g. '(' with ')', '[' with ']', '{' with '}'
            if (char == ')' and last != '(') or \
               (char == ']' and last != '[') or \
               (char == '}' and last != '{'):
                return False
    # valid iff no unmatched openers remain, i.e. empty stack -> valid
    return not stack
```

<!-- grade_id: correct_is_valid, points: 5, grade: true, solution: false -->
```python
assert is_valid('()') == True
assert is_valid('([{]}') == False
assert is_valid('([{}])') == True
```